# Homework 2 (Total Points: 10) <a class="anchor" id="top"></a>


## Submission instructions:

- The assignment is designed for teams of two students.
- Add your solutions to the cells with the following comment `# YOUR CODE HERE`, which denotes that these sections are graded.
- The notebook that you will submit has to have the student ids, separated by underscores (E.g., `12341234_12341234.ipynb`). 
- This will be parsed by a regexp, so please double check your filename.
- Only one member of each group should submit the notebook file to canvas.
- Make sure to check that your notebook runs before submission. A quick way to do this is to restart the kernel and run all the cells.  
- Please do not delete cells. Removing cells can lead to grade deduction. You are allowed to add cells for performing tests or adding auxiliary methods in the given cells.
- Note, that you are not allowed to submit notebooks created with Google Colab. (The reason is that such notebooks do not work with our semi-automated grading system. You may work with Google Colab on developing your solutions. Please then just copy your answers into a new notebook that you haven't opened with Google Colab.)
- Add your name and e-mail in the following cell:

```python
## Add your information in here.
NAME = "Giulio Starace"
NAME2 = "Matteo Rosati"
EMAIL = "giulio.starace@student.uva.nl"
EMAIL2 = "matteo.rosati@student.uva.nl"
```

**Assignment instructions Goals**:

In this assignment, you will program several algorithms related to BombKnightSudoku (a variant of Sudoku) on a $(k*k)^2$ size grid, for varying sizes of $k$. You will use several problem solving and search methods.

The homework assignment consists of several parts (a)–(d). You get 1 point for free. For the rest parts you need to develop the following algorithms for solving sudoku game:

 - Part (a) 4pts: solving puzzles using SAT
 - Part (b) 4pts: solving puzzles using ASP
 - Part (c) 2pts: generating puzzles with unique solutions using ASP (or you may opt to solve an easier variant for a maximum of 1 point)

**Game information: BombKnightSudoku's**

Here follows a description of inputs and solutions for **BombKnightSudoku's**. These notions are similar (but not exactly the same) as for the Sudoku's that you might have seen elsewhere, generalized to arbitrary values of k (typically, $k=3$).

An input consists of a $k$-by-$k$ grid of blocks, where each consisting of a $k$-by-$k$ grid of cells, forming a $(k*k)$-by-$(k*k)$ grid of cells. Each cell can (but must not) contain a value between 1 and $k*k$.

An example is the following grid, for $k$=3, written out as $k*k$ lines, each consisting of $k*k$ numbers, separated by spaces. Here the value $0$ represents that a cell does not contain a value between $1$ and $k*k$.

```
0 0 0 0 0 0 0 0 0
0 3 0 0 8 0 0 0 0
0 0 0 0 0 0 0 0 0
0 0 0 0 4 0 0 0 5
0 0 0 0 0 0 0 2 7
0 0 0 0 0 0 3 0 0
0 0 0 0 0 0 9 0 0
0 0 0 0 5 6 0 0 0
0 0 0 0 0 0 0 0 0
```

Here is the same input, where now the $k$-by-$k$ blocks are separated from each other with whitespace.

```
0 0 0  0 0 0  0 0 0
0 3 0  0 8 0  0 0 0
0 0 0  0 0 0  0 0 0

0 0 0  0 4 0  0 0 5
0 0 0  0 0 0  0 2 7
0 0 0  0 0 0  3 0 0

0 0 0  0 0 0  9 0 0
0 0 0  0 5 6  0 0 0
0 0 0  0 0 0  0 0 0
```

A solution for a sudoku input is a $(k*k)$-by-$(k*k)$ grid (for the same $k$), where:

 - Each cell contains a value between 1 and $k*k$.
If a cell ($i$,$j$) contains a value u in the input, then the cell ($i$,$j$) in the solution must contain the same value u.
 - Each two different cells in the same row must contain different values.
 - Each two different cells in the same column must contain different values.
 - Each two different cells in the same $k*k$ block must contain different values.
 - Each two different cells that are directly adjacent (horizontally, diagonally or vertically) must contain different values. In other words, a cell cannot contain the same value as any of the eight cells surrounding it.
 - Each two different cells that can be reached from each other with a single knight's move (as in Chess) must contain different values.

(The last two constraints are what distinguish BombKnightSudoku's from regular Sudoku's.)

For example, a solution for the input mentioned above is the following:

```
5 6 2 7 3 4 1 8 9
7 3 4 1 8 9 5 6 2
1 8 9 5 6 2 7 3 4
6 2 7 3 4 1 8 9 5
3 4 1 8 9 5 6 2 7
8 9 5 6 2 7 3 4 1
2 7 3 4 1 8 9 5 6
4 1 8 9 5 6 2 7 3
9 5 6 2 7 3 4 1 8
```

**Representation of inputs and outputs**

Inputs are represented as a list of lists of numbers.
For example, the BombKnightSudoku input above is represented as:

```
input =
    [[0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 3, 0, 0, 8, 0, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 0, 0, 0, 4, 0, 0, 0, 5],
     [0, 0, 0, 0, 0, 0, 0, 2, 7],
     [0, 0, 0, 0, 0, 0, 3, 0, 0],
     [0, 0, 0, 0, 0, 0, 9, 0, 0],
     [0, 0, 0, 0, 5, 6, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0]]
```

Solutions are represented similarly, for example:

```
solution =
    [[5, 6, 2, 7, 3, 4, 1, 8, 9],
     [7, 3, 4, 1, 8, 9, 5, 6, 2],
     [1, 8, 9, 5, 6, 2, 7, 3, 4],
     [6, 2, 7, 3, 4, 1, 8, 9, 5],
     [3, 4, 1, 8, 9, 5, 6, 2, 7],
     [8, 9, 5, 6, 2, 7, 3, 4, 1],
     [2, 7, 3, 4, 1, 8, 9, 5, 6],
     [4, 1, 8, 9, 5, 6, 2, 7, 3],
     [9, 5, 6, 2, 7, 3, 4, 1, 8]]
```

Note that some test_inputs are given in the auxiliary python file `sudoku_func.py` which are used for testing your methods.

The necessary libraries for developing all necessary code can be found in the following block:

```python
# Necessary libraries that will be employed
import itertools
from copy import deepcopy
from dataclasses import dataclass
from typing import List

from sudoku_func import *

#!{sys.executable} -m pip install clingo
#!{sys.executable} -m pip install python-sat
import clingo
import pysat
from pysat.formula import CNF
from pysat.solvers import MinisatGH
```

**N.B.** There are two *pysat* libraries. Please make sure to pip install python-sat (*not* pysat).

**N.B.** Please make sure *not* to remove or change the code you have been given (e.g. do not add, modify or remove the function names, names of the parameters, etc.).

## (a): solving puzzles using SAT

Again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into a propositional CNF formula &varphi;, calling a SAT solver to find satisfying assignments for &varphi;, and translating these assignments to solutions for the BombKnightSudoku input.

```python
# You are allowed to add cells for [...] adding auxiliary methods in the given cells.


def assignment_to_sudoku(assignment, vpool, max_value, fillable_sudoku=None):
    """
    Converts a pysat assignment to a sudoku array

    Parameters:
        assignment (list(int)): pysat CNF assignment
        vpool (pysat.formula.IDPool): pysat IDPool instance to be used for mapping
            integer variables back to objects
        max_value (int): maximum value that can be in our sudoku. k**2.
        fillable_sudoku (list(list(int))): prepared list ready to be filled
            with solution, optional
    Returns:
        The resulting sudoku solution (list(list(int)))
    """
    if fillable_sudoku is None:
        solution = [[0 for _ in range(max_value)] for _ in range(max_value)]
    else:
        solution = deepcopy(fillable_sudoku)
    for variable in assignment:
        if variable > 0:
            (i, j, val) = vpool.obj(variable)
            solution[i][j] = val
    return solution


def ensure_diff_to_target(target, group, cnf, vpool, min_value, max_value):
    """
    Ensures that each cell in the group has a different value from the target cell

    This is slightly different from ensure_group_set, as we only care about the
    difference in value from the target cell.

    Parameters:
        target (tuple): (row, column) of target cell
        group list(tuple): (row, column) of each cell in group
        cnf (pysat.formula.CNF): instance of pysat CNF formula,
            to be modified (appended to) in place
        vpool (pysat.formula.IDPool): pysat IDPool instance to be used for mapping
            integer variables back to objects
        min_value (int): the minimum value that can be taken in the sudoku
            (usually 1)
        max_value (int): the maximum value that can be taken in the sudoku
            (usually k**2)

    Updates the passed cnf with new clauses in place.
    """
    target_i, target_j = target
    for value in range(min_value, max_value + 1):
        for i, j in group:
            cnf.append(
                [
                    -1 * vpool.id((target_i, target_j, value)),
                    -1 * vpool.id((i, j, value)),
                ]
            )


def ensure_group_set(cells_by_group, cnf, vpool, min_value, max_value):
    """
    Ensures each group of cells in the cell_by_group array
    is a set such that each cell has a unique value within that group.

    an example of a cell_group could be a given row, s.t. cells_by_group is then
    a list where each element is the list of cells in a given row

    Parameters:
        cells_by_group list(list(tuple)): list of groups, where each element (group)
            is a list of the cells ((row, column) tuple) for that particular group
        cnf (pysat.formula.CNF): instance of pysat CNF formula,
            to be modified (appended to) in place
        vpool (pysat.formula.IDPool): pysat IDPool instance to be used for mapping
            integer variables back to objects
        min_value (int): the minimum value that can be taken in the sudoku
            (usually 1)
        max_value (int): the maximum value that can be taken in the sudoku
            (usually k**2)

    updates the passed cnf with new clauses in place.
    """
    for cell_group in cells_by_group:
        for value in range(min_value, max_value + 1):
            for cell_pair in itertools.combinations(cell_group, 2):
                cnf.append([-1 * vpool.id((i, j, value)) for (i, j) in cell_pair])


def find_adj_diag(cell, n, include_cell=False):
    """
    finds adjacent diagonal cells of a given cell in n*n matrix

    Parameters:
        cell (tuple): (row, column) tuple
        n (int): the size of the n*n matrix the cell is in
            necessary to check boundary conditions
        include_cell (bool): whether to include the cell in the result,
            default False

    Returns:
        (list(tuple)) List of (row, column) tuples corr. to adjacent diagonal cells
    """
    row, col = cell
    adj_diag = [
        (row + i, col + j)
        for i in (-1, 1)
        for j in (-1, 1)
        if 0 <= row + i < n and 0 <= col + j < n
    ]
    if include_cell:
        adj_diag.append(cell)
    return adj_diag


def find_knights(cell, n, include_cell=False):
    """
    finds knight's move cells of a given cell in n*n matrix

    Parameters:
        cell (tuple): (row, column) tuple
        n (int): the size of the n*n matrix the cell is in
            necessary to check boundary conditions
        include_cell (bool): whether to include the cell in the result,
            default False

    Returns:
        (list(tuple)) List of (row, column) tuples corr. to knights
    """
    row, col = cell
    knights = [
        (row + i, col + j)
        for i in [-2, -1, 1, 2]
        for j in [-2, -1, 1, 2]
        if 0 <= row + i < n and 0 <= col + j < n and abs(j) != abs(i)
    ]
    if include_cell:
        knights.append(cell)
    return knights
```

<!-- grade_id: cell-940a32cfcc687489, points: 0, grade: false, solution: true -->
```python
def solve_sudoku_SAT(sudoku: List[List[int]], k: int):
    """
    Solve a sudoku input by encoding the problem into SAT,
    calling a SAT solver, and retrieving the solution for
    the sudoku input from a satisfying truth assignment given
    by the SAT solver.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """
    # YOUR CODE HERE
    min_value = 1
    max_value = int(k**2)
    # https://krr-talk.computational-complexity.nl/t/idpool-for-question-1/235/2
    vpool = pysat.formula.IDPool()  # we use tuples of (row, col, value)
    cnf = CNF()
    # {{{ -------------------- Useful precomputations --------------------------
    cells = list(itertools.product(range(max_value), repeat=2))
    non_border_cells = [
        cell for cell in cells if (0 not in cell) and (max_value - 1 not in cell)
    ]
    centers_idxs = range(max_value)[1::k]
    # optmztn: only need to check bomb cnstrnt on these cells (given other cnstrnts)
    non_block_center_cells = [
        cell
        for cell in non_border_cells
        if cell[0] not in centers_idxs or cell[1] not in centers_idxs
    ]
    # now, organize cells by groupings (by row, by col, by block, etc)
    cells_by_row = [[cell for cell in cells if cell[0] == j] for j in range(max_value)]
    cells_by_col = [[cell for cell in cells if cell[1] == j] for j in range(max_value)]
    cells_by_block = [
        [
            cell
            for cell in cells
            if cell[0] in range(j * k, j * k + k) and cell[1] in range(i * k, i * k + k)
        ]
        for j in range(k)
        for i in range(k)
    ]
    # given other cnstrts, only need to check adjacent diagonals for bomb constraint
    cells_by_adj_diag = {
        cell: find_adj_diag(cell, max_value, False) for cell in non_block_center_cells
    }
    cells_by_knights = {cell: find_knights(cell, max_value, False) for cell in cells}
    # }}}
    # {{{---------------------- encoding the problem into SAT ---------------------
    # Each cell contains a value between 1 and k*k
    for (i, j) in cells:
        input_value = sudoku[i][j]
        # If a cell (i,j) contains a value u in the input
        if input_value != 0:
            # then the cell (i, j) in the solution must contain the same value u
            cnf.append([vpool.id((i, j, input_value))])
        else:
            # each (remaining) cell contains a value between 1 and k**2
            cnf.append(
                [vpool.id((i, j, val)) for val in range(min_value, max_value + 1)]
            )
    # Each two different cells in the same row must contain different values.
    ensure_group_set(cells_by_row, cnf, vpool, min_value, max_value)
    # Each two different cells in the same column must contain different values.
    ensure_group_set(cells_by_col, cnf, vpool, min_value, max_value)
    # Each two different cells in the same k*k block must contain different values
    ensure_group_set(cells_by_block, cnf, vpool, min_value, max_value)
    # Each two different cells that are directly adjacent must contain diff values
    for cell, adj_cells in cells_by_adj_diag.items():
        ensure_diff_to_target(cell, adj_cells, cnf, vpool, min_value, max_value)
    # Each two diff cells that can be reached by knights move must contain diff values
    for cell, knight_cells in cells_by_knights.items():
        ensure_diff_to_target(cell, knight_cells, cnf, vpool, min_value, max_value)
    #  }}}
    #  {{{------ calling a SAT solver and converting assingment to sudoku----------
    with MinisatGH(bootstrap_with=cnf.clauses) as m:
        solvable = True
        while solvable:
            solvable = m.solve()
            if not solvable:
                # there are no remaining assignments
                return
                yield  # yield empty generator
            else:
                assignment = m.get_model()
                solution = assignment_to_sudoku(assignment, vpool, max_value, sudoku)
                # add negated assignment as a disjunction
                m.add_clause([-1 * var for var in assignment])
                yield solution
    #  }}}
```

The following code will test whether your provided solution has the correct number of solution for the test puzzles.

<!-- grade_id: cell-4f0c31f4bfc245aa, points: 4, grade: true, solution: false -->
```python
# This cell will test whether your code for solve_sudoku_SAT()
# finds the right number of solutions on several test inputs.

from sudoku_func import check_num_solutions
from sudoku_func import test_inputs

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert (
        check_num_solutions(
            test_input.sudoku, test_input.k, test_input.num_solutions, solve_sudoku_SAT
        )
        == True
    )
    print("  Found exactly the set of correct solutions")

# (There will also be some hidden tests..)
```

## (b): solving puzzles using ASP

Yet again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to solutions for the BombKnightSudoku input.

```python
def parse_sudoku(sudoku: List[List[int]]):
    """
    Parses sudoku starting state and generates ASP fact string

    Parameters:
        sudoku (list(list(int))): An input puzzle.

    Returns:
        A string of ASP input strings.
    """
    parsed = ""
    if not sudoku:
        return parsed
    for nr, r in enumerate(sudoku):
        for nc, val in enumerate(r):
            if val != 0:
                new_fact = f"place({nr},{nc},{val}).\n"
                parsed += new_fact
    return parsed


def program_string(sudoku: List[List[int]], k: int):
    """
    Returns the ASP program string.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.
    """
    # preamble string
    preamble = (
        f"#const k={k}.\n"
        + "row(0..k**2-1).\n"
        + "col(0..k**2-1).\n"
        + "val(1..k**2).\n"
    )

    # parse the sudoku input
    start_state = parse_sudoku(sudoku)

    # build string
    # generate candidate solutions
    generate = "1 {place(R, C, V) : val(V)} 1 :- row(R), col(C).\n"
    helper = "valid(R, C) :- row(R), col(C).\n"

    initial = preamble + start_state + generate + helper

    # define blocks and what it means for a square to be reachable by knight or king move
    define = (
        "in_block(R, C, SR, SC) :- valid(SR, SC), valid(R, C),\n"
        + "     SR \ k == 0, SC \ k = 0, R - SR < k, C - SC < k, R - SR >= 0, C - SC >= 0.\n"
        + "kreachable(R, C, AR, AC) :- valid(R, C), valid(AR, AC),\n"
        + "     |C - AC| = 1, |R - AR| <= 1.\n"
        + "kreachable(R, C, AR, AC) :- valid(R, C), valid(AR, AC),\n"
        + "     |R - AR| = 1, |C - AC| <= 1.\n"
        + "nreachable(R, C, AR, AC) :- valid(R, C), valid(AR, AC),\n"
        + "     |R - AR| == 1, |C - AC| == 2.\n"
        + "nreachable(R, C, AR, AC) :- valid(R, C), valid(AR, AC), \n"
        + "     |R - AR| == 2, |C - AC| == 1.\n"
        + "reachable(R, C, AR, AC) :- nreachable(AR, AC, R, C).\n"
        + "reachable(R, C, AR, AC) :- kreachable(AR, AC, R, C).\n"
    )

    # add restrictions (no same values on same row or column, no same king or knight adjacent values)
    restrict = (
        ":- place(R1, C, V), place(R2, C, V), R1 != R2.\n"
        + ":- place(R, C1, V), place(R, C2, V), C1 != C2.\n"
        + ":- place(R, C, V), place(R1, C1, V),\n"
        + "    in_block(R, C, SC, EC),  in_block(R1, C1, SC, EC), R != R1, C != C1.\n"
        + ":- place(R, C, V), place(R1, C1, V), reachable(R, C, R1, C1).\n"
        + f"#show place/3."
    )

    return initial + define + restrict


def model_parse(model: clingo.solving.Model, k: int):
    """
    Parses a model back to sudoku format (list(list(int))).
    """
    side = k**2
    solution = [[0 for _ in range(side)] for _ in range(side)]
    for atom in model.symbols(shown=True):
        if atom.name == "place":
            r = atom.arguments[0].number
            c = atom.arguments[1].number
            v = atom.arguments[2].number
            solution[r][c] = v
    return solution
```

<!-- grade_id: cell-e08fbc765eba6e6d, points: 0, grade: false, solution: true -->
```python
def solve_sudoku_ASP(sudoku, k):
    """
    Solve a sudoku input by encoding the problem into ASP, 
    calling an ASP solver, and retrieving the solution 
    for the sudoku input from an answer set given by the ASP solver.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # build program string
    program = program_string(sudoku, k)

    # Load the answer set program, and call the grounder
    control = clingo.Control()
    control.add("base", [], program)
    control.ground([("base", [])])

    # Ask clingo to find all models (using an upper bound of 0 gives all models)
    control.configuration.solve.models = 0

    # Call the clingo solver, passing on the function on_model for when an answer set is found
    # with control.solve(on_model = on_model) as solutions:
    for i in control.solve(yield_=True):
        yield model_parse(i, k)

    # Print a message when no answer set was found
    # if answer.satisfiable == False:
    #     print("No answer sets")
```

The following code will test whether your provided solution has the correct number of solution for the test puzzles.

<!-- grade_id: cell-d58b222e808c0f0e, points: 4, grade: true, solution: false -->
```python
# This cell will test whether your code for solve_sudoku_ASP()
# finds the right number of solutions on several test inputs.

from sudoku_func import check_num_solutions
from sudoku_func import test_inputs

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert (
        check_num_solutions(
            test_input.sudoku, test_input.k, test_input.num_solutions, solve_sudoku_ASP
        )
        == True
    )
    print("  Found exactly the set of correct solutions")

# (There will also be some hidden tests..)
```

## (c): generating puzzles with unique solutions using ASP

Implement an algorithm that generates `9`-by-`9` BombKnightSudoku inputs (i.e., `k=3`) that:
1. have exactly 10 cells filled with a number between `1` and `9`,
1. and have a unique solution.

That is, your algorithm should output a generator that iterates over BombKnightSudoku inputs with the required properties.

Your algorithm should use the approach of translating this problem to an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to BombKnightSudoku inputs with the required properties.

(To be clear: you may not 'hardcode' a list of BombKnightSudoku inputs in the answer set program *P*.)

There are two variants of this assignment: a harder variant (d.I) and an easier variant (d.II). You may choose which variant you would like to solve.
If you opt for the harder variant (d.I), you can get a maximum of 2 points for this assignment.
If you opt for the easier variant (d.II), you can get a maximum of 1 points for this assignment.


### (c.I) Harder variant (2pts):

To solve the harder variant, your answer set program *P* should be such that for **each** BombKnightSudoku input *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you should express in *P* the constraints 1. (that exactly 10 cells are filled) and 2. (that the input has a unique solution).

*Hint:* Use the technique of saturation to express constraint 2.

### (c.II) Easier variant (1pt):

To solve the easier variant, your answer set program *P* only needs to be such that for **a subset of** BombKnightSudoku inputs *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you could express in *P*, in addition to the constraint 1. (that exactly 10 cells are filled), some other property that entails constraint 2. (that the input has a unique solution).

*Hint:* If a Sudoku input can be solved using some [basic strategies](http://pi.math.cornell.edu/~mec/Summer2009/Mahmood/Solve.html), it must have a unique solution. Consider variants of these basic strategies for BombKnightSudoku.

```python
def saturation_string(k: int, p: int):
    """
    Returns the ASP saturation program string.

    Parameters:
        k (int): The dimension of the sudoku input.
        p (int): How many numbers to place in problem.
    """

    # sudoku generate solutions without any starting set up
    head = f"#const p = {p}.\n" + program_string([], k)
    head = "\n".join(head.split("\n")[:-1])

    # hide some cells from those placed on the board (or fill with zeros)
    hide = (
        "{ hide(R, C) : valid(R, C) }.\n"
        + "on_board(R, C, V) :- not hide(R, C), place(R, C, V).\n"
    )

    # ensure exactly p cells are filled in in our problem
    exactly_p = ":- not p { on_board(R, C, V) : valid(R, C), val(V) } p."

    program = head + hide + exactly_p

    
    # generate candidate new solutions,
    # express cases in which new solution is incorrect:
    # - new solution has different values placed on board than filled in prior
    # - new soution is the same as old solution
    # - new solution breaks rules of sudoku
    saturation = (
        "new_place(R, C, V) : val(V) :- valid(R, C).\n"
        + "incorrect :- on_board(R, C, V), new_place(R, C, W), V != W.\n"
        + "incorrect :- new_place(R, C, V) : place(R, C, V).\n"
        + "incorrect :- new_place(R1, C, V), new_place(R2, C, V), R1 != R2,\n"
        + "    valid(R1, C), valid(R2, C).\n"
        + "incorrect :- new_place(R, C1, V), new_place(R, C2, V), C1 != C2,\n"
        + "    valid(R, C1), valid(R, C2).\n"
        + "incorrect :- new_place(R, C, V), new_place(R1, C1, V),\n"
        + "    in_block(R, C, SC, EC),  in_block(R1, C1, SC, EC),\n"
        + "    R != R1, C != C1, valid(R, C), valid(R1, C1).\n"
        + "incorrect :- new_place(R, C, V), new_place(R1, C1, V), reachable(R, C, R1, C1),\n"
        + "    valid(R, C), valid(R1, C1).\n"
    )
    program += saturation
        
    # finish saturation script, discard unsaturated   
    program += (
        "saturate :- incorrect.\n"
        + "incorrect :- saturate.\n"
        + "new_place(R, C, V) :- val(V), valid(R, C), saturate.\n"
        + ":- not saturate.\n"
    )

    program += saturation

    program += "#show on_board/3."

    return program


def sat_model_parse(model: clingo.solving.Model, k: int):
    """
    Parses a saturation input back to sudoku, format (list(list(int))).
    """
    side = k**2
    board = [[0 for _ in range(side)] for _ in range(side)]
    for atom in model.symbols(shown=True):
        if atom.name == "on_board":
            r = atom.arguments[0].number
            c = atom.arguments[1].number
            v = atom.arguments[2].number
            board[r][c] = v
    return board
```

<!-- grade_id: cell-fb57584e65e58ba1, points: 0, grade: false, solution: true -->
```python
def generate_puzzles_with_unique_solutions():
    """
    Algorithm that generates BombKnightSudoku puzzles with a unique solution
    and where exactly 10 cells are filled.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    k = 3
    p = 10

    program = saturation_string(k, p)

    # Load the answer set program, and call the grounder
    control = clingo.Control()
    control.add("base", [], program)
    control.ground([("base", [])])

    # Ask clingo to find all models (using an upper bound of 0 gives all models)
    control.configuration.solve.models = 0

    # Call the clingo solver, passing on the function on_model for when an answer set is found
    # with control.solve(on_model = on_model) as solutions:
    for i in control.solve(yield_=True):
        yield sat_model_parse(i, k)
```

<!-- grade_id: cell-d86460df48d77382, points: 2, grade: true, solution: false -->
```python
# This cell will test whether your function yields
# (at least) 10 different sudoku puzzles with unique solutions.

# (These tests are hidden.)
```