# Homework 2 (Total Points: 10) <a class="anchor" id="top"></a>


## Submission instructions:

- The assignment is designed for teams of two students.
- Add your solutions to the cells with the following comment `# YOUR CODE HERE`, which denotes that these sections are graded.
- The notebook that you will submit has to have the student ids, separated by underscores (E.g., `12341234_12341234.ipynb`). 
- This will be parsed by a regexp, so please double check your filename.
- Only one member of each group should submit the notebook file to canvas.
- Make sure to check that your notebook runs before submission. A quick way to do this is to restart the kernel and run all the cells.  
- Please do not delete cells. Removing cells can lead to grade deduction. You are allowed to add cells for performing tests or adding auxiliary methods in the given cells.
- Note, that you are not allowed to submit notebooks created with Google Colab. (The reason is that such notebooks do not work with our semi-automated grading system. You may work with Google Colab on developing your solutions. Please then just copy your answers into a new notebook that you haven't opened with Google Colab.)
- Add your name and e-mail in the following cell:

```python
## Add your information in here.
NAME = ""
NAME2 = ""
EMAIL = ""
EMAIL2 = ""
```

**Assignment instructions Goals**:

In this assignment, you will program several algorithms related to BombKnightSudoku (a variant of Sudoku) on a $(k*k)^2$ size grid, for varying sizes of $k$. You will use several problem solving and search methods.

The homework assignment consists of several parts (a)–(d). You get 1 point for free. For the rest parts you need to develop the following algorithms for solving sudoku game:

 - Part (a) 4pts: solving puzzles using SAT
 - Part (b) 4pts: solving puzzles using ASP
 - Part (c) 2pts: generating puzzles with unique solutions using ASP (or you may opt to solve an easier variant for a maximum of 1 point)

**Game information: BombKnightSudoku's**

Here follows a description of inputs and solutions for **BombKnightSudoku's**. These notions are similar (but not exactly the same) as for the Sudoku's that you might have seen elsewhere, generalized to arbitrary values of k (typically, $k=3$).

An input consists of a $k$-by-$k$ grid of blocks, where each consisting of a $k$-by-$k$ grid of cells, forming a $(k*k)$-by-$(k*k)$ grid of cells. Each cell can (but must not) contain a value between 1 and $k*k$.

An example is the following grid, for $k$=3, written out as $k*k$ lines, each consisting of $k*k$ numbers, separated by spaces. Here the value $0$ represents that a cell does not contain a value between $1$ and $k*k$.

```
0 0 0 0 0 0 0 0 0
0 3 0 0 8 0 0 0 0
0 0 0 0 0 0 0 0 0
0 0 0 0 4 0 0 0 5
0 0 0 0 0 0 0 2 7
0 0 0 0 0 0 3 0 0
0 0 0 0 0 0 9 0 0
0 0 0 0 5 6 0 0 0
0 0 0 0 0 0 0 0 0
```

Here is the same input, where now the $k$-by-$k$ blocks are separated from each other with whitespace.

```
0 0 0  0 0 0  0 0 0
0 3 0  0 8 0  0 0 0
0 0 0  0 0 0  0 0 0

0 0 0  0 4 0  0 0 5
0 0 0  0 0 0  0 2 7
0 0 0  0 0 0  3 0 0

0 0 0  0 0 0  9 0 0
0 0 0  0 5 6  0 0 0
0 0 0  0 0 0  0 0 0
```

A solution for a sudoku input is a $(k*k)$-by-$(k*k)$ grid (for the same $k$), where:

 - Each cell contains a value between 1 and $k*k$.
If a cell ($i$,$j$) contains a value u in the input, then the cell ($i$,$j$) in the solution must contain the same value u.
 - Each two different cells in the same row must contain different values.
 - Each two different cells in the same column must contain different values.
 - Each two different cells in the same $k*k$ block must contain different values.
 - Each two different cells that are directly adjacent (horizontally, diagonally or vertically) must contain different values. In other words, a cell cannot contain the same value as any of the eight cells surrounding it.
 - Each two different cells that can be reached from each other with a single knight's move (as in Chess) must contain different values.

(The last two constraints are what distinguish BombKnightSudoku's from regular Sudoku's.)

For example, a solution for the input mentioned above is the following:

```
5 6 2 7 3 4 1 8 9
7 3 4 1 8 9 5 6 2
1 8 9 5 6 2 7 3 4
6 2 7 3 4 1 8 9 5
3 4 1 8 9 5 6 2 7
8 9 5 6 2 7 3 4 1
2 7 3 4 1 8 9 5 6
4 1 8 9 5 6 2 7 3
9 5 6 2 7 3 4 1 8
```

**Representation of inputs and outputs**

Inputs are represented as a list of lists of numbers.
For example, the BombKnightSudoku input above is represented as:

```
input =
    [[0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 3, 0, 0, 8, 0, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 0, 0, 0, 4, 0, 0, 0, 5],
     [0, 0, 0, 0, 0, 0, 0, 2, 7],
     [0, 0, 0, 0, 0, 0, 3, 0, 0],
     [0, 0, 0, 0, 0, 0, 9, 0, 0],
     [0, 0, 0, 0, 5, 6, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0]]
```

Solutions are represented similarly, for example:

```
solution =
    [[5, 6, 2, 7, 3, 4, 1, 8, 9],
     [7, 3, 4, 1, 8, 9, 5, 6, 2],
     [1, 8, 9, 5, 6, 2, 7, 3, 4],
     [6, 2, 7, 3, 4, 1, 8, 9, 5],
     [3, 4, 1, 8, 9, 5, 6, 2, 7],
     [8, 9, 5, 6, 2, 7, 3, 4, 1],
     [2, 7, 3, 4, 1, 8, 9, 5, 6],
     [4, 1, 8, 9, 5, 6, 2, 7, 3],
     [9, 5, 6, 2, 7, 3, 4, 1, 8]]
```

Note that some test_inputs are given in the auxiliary python file `sudoku_func.py` which are used for testing your methods.

The necessary libraries for developing all necessary code can be found in the following block:

```python
# Necessary libraries that will be employed
import itertools
from copy import deepcopy
from dataclasses import dataclass
from typing import List

from sudoku_func import *

#!{sys.executable} -m pip install clingo
#!{sys.executable} -m pip install python-sat
import clingo
import pysat
from pysat.formula import CNF
from pysat.solvers import MinisatGH
```

**N.B.** There are two *pysat* libraries. Please make sure to pip install python-sat (*not* pysat).

**N.B.** Please make sure *not* to remove or change the code you have been given (e.g. do not add, modify or remove the function names, names of the parameters, etc.).

## (a): solving puzzles using SAT

Again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into a propositional CNF formula &varphi;, calling a SAT solver to find satisfying assignments for &varphi;, and translating these assignments to solutions for the BombKnightSudoku input.

<!-- grade_id: cell-940a32cfcc687489, points: 0, grade: false, solution: true -->
```python
def solve_sudoku_SAT(sudoku, k):
    """
    Solve a sudoku input by encoding the problem into SAT, calling a SAT solver, and retrieving the solution for the sudoku input from a satisfying truth assignment given by the SAT solver.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

The following code will test whether your provided solution has the correct number of solution for the test puzzles.

<!-- grade_id: cell-4f0c31f4bfc245aa, points: 4, grade: true, solution: false -->
```python
# This cell will test whether your code for solve_sudoku_SAT()
# finds the right number of solutions on several test inputs.

from sudoku_func import check_num_solutions
from sudoku_func import test_inputs

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(
        test_input.sudoku,
        test_input.k,
        test_input.num_solutions,
        solve_sudoku_SAT
    ) == True
    print("  Found exactly the set of correct solutions")

# (There will also be some hidden tests..)

```

## (b): solving puzzles using ASP

Yet again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to solutions for the BombKnightSudoku input.

<!-- grade_id: cell-e08fbc765eba6e6d, points: 0, grade: false, solution: true -->
```python
def solve_sudoku_ASP(sudoku, k):
    """
    Solve a sudoku input by encoding the problem into ASP, calling an ASP solver, and retrieving the solution for the sudoku input from an answer set given by the ASP solver.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

The following code will test whether your provided solution has the correct number of solution for the test puzzles.

<!-- grade_id: cell-d58b222e808c0f0e, points: 4, grade: true, solution: false -->
```python
# This cell will test whether your code for solve_sudoku_ASP()
# finds the right number of solutions on several test inputs.

from sudoku_func import check_num_solutions
from sudoku_func import test_inputs

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(
        test_input.sudoku,
        test_input.k,
        test_input.num_solutions,
        solve_sudoku_ASP
    ) == True
    print("  Found exactly the set of correct solutions")

# (There will also be some hidden tests..)

```

## (c): generating puzzles with unique solutions using ASP

Implement an algorithm that generates `9`-by-`9` BombKnightSudoku inputs (i.e., `k=3`) that:
1. have exactly 10 cells filled with a number between `1` and `9`,
1. and have a unique solution.

That is, your algorithm should output a generator that iterates over BombKnightSudoku inputs with the required properties.

Your algorithm should use the approach of translating this problem to an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to BombKnightSudoku inputs with the required properties.

(To be clear: you may not 'hardcode' a list of BombKnightSudoku inputs in the answer set program *P*.)

There are two variants of this assignment: a harder variant (d.I) and an easier variant (d.II). You may choose which variant you would like to solve.
If you opt for the harder variant (d.I), you can get a maximum of 2 points for this assignment.
If you opt for the easier variant (d.II), you can get a maximum of 1 points for this assignment.


### (c.I) Harder variant (2pts):

To solve the harder variant, your answer set program *P* should be such that for **each** BombKnightSudoku input *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you should express in *P* the constraints 1. (that exactly 10 cells are filled) and 2. (that the input has a unique solution).

*Hint:* Use the technique of saturation to express constraint 2.

### (c.II) Easier variant (1pt):

To solve the easier variant, your answer set program *P* only needs to be such that for **a subset of** BombKnightSudoku inputs *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you could express in *P*, in addition to the constraint 1. (that exactly 10 cells are filled), some other property that entails constraint 2. (that the input has a unique solution).

*Hint:* If a Sudoku input can be solved using some [basic strategies](http://pi.math.cornell.edu/~mec/Summer2009/Mahmood/Solve.html), it must have a unique solution. Consider variants of these basic strategies for BombKnightSudoku.

<!-- grade_id: cell-fb57584e65e58ba1, points: 0, grade: false, solution: true -->
```python
def generate_puzzles_with_unique_solutions():
    """
    Algorithm that generates BombKnightSudoku puzzles with a unique solution
    and where exactly 10 cells are filled.
    
    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d86460df48d77382, points: 2, grade: true, solution: false -->
```python
# This cell will test whether your function yields
# (at least) 10 different sudoku puzzles with unique solutions.

# (These tests are hidden.)

```